<!DOCTYPE html>
<html>
<head>
   <title>CSS Animations - initAnimationEvent called twice</title>
   <link rel="author" title="Opera Software ASA" href="http://www.opera.com/">
   <link rel="help" href="https://drafts.csswg.org/css3-animations/#animation-events-">
   <meta name="flags" content="dom">
   <meta name="assert" content="If initAnimationEvent is called multiple times
   before being dispatched, the final invocation takes precedence">
   <script src="/resources/testharness.js"></script>
   <script src="/resources/testharnessreport.js"></script>
   <script type="text/javascript">
      setup({explicit_done:true});
      onload = function() {
         var testName = document.querySelector('title').firstChild.data;
         var log = document.querySelectorAll('p')[0].firstChild;
         var passed;
         try {
            var evt = document.createEvent('AnimationEvent');
            evt.initAnimationEvent('animationend', true, false, 'Before', 0);
            evt.initAnimationEvent('animationend', true, false, 'After', 1);
            passed = (evt.animationName == 'After' && evt.elapsedTime == 1);
         } catch(e) {
            passed = false;
         }
         test(function(){assert_true(passed);}, testName); done();
         log.data = (passed ? 'PASS' : 'FAIL');
      };
   </script>
   <style>iframe{display:none}</style>
</head>
<body>
   <p>Testing... (JavaScript required)</p>
</body>
</html>
