<!DOCTYPE html>
<html>
<head>
   <title>CSS Animations - typeArg in initAnimationEvent</title>
   <link rel="author" title="Opera Software ASA" href="http://www.opera.com/">
   <link rel="help" href="https://drafts.csswg.org/css3-animations/#animation-events-">
   <meta name="flags" content="dom">
   <meta name="assert" content="typeArg in initAnimationEvent specifies the
   event type">
   <script src="/resources/testharness.js"></script>
   <script src="/resources/testharnessreport.js"></script>
   <script type="text/javascript">
      setup({explicit_done:true});
      onload = function() {
         var testName = document.querySelector('title').firstChild.data;
         var div = document.querySelector('div');
         var log = document.querySelectorAll('p')[0].firstChild;
         var passed = undefined;
         div.addEventListener('animationstart', function() {
            passed = false;
         });
         div.addEventListener('animationend', function() {
            passed = (passed !== false);
         });
         setTimeout(function() {
            test(function(){assert_true(passed);}, testName); done();
            log.data = (passed ? 'PASS' : 'FAIL');
         }, 500);
         try {
            var evt = document.createEvent('AnimationEvent');
            evt.initAnimationEvent('animationend', true, false, 'test', 0);
            div.dispatchEvent(evt);
         } catch(e) {
            passed = false;
         }
      };
   </script>
   <style>iframe{display:none}</style>
</head>
<body>
   <p>Testing... (JavaScript required)</p>
   <div></div>
</body>
</html>
