// Copyright (C) 2015 André Bargull. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-typedarray-constructors
description: >
  Float32Array is a constructor function.
features: [TypedArray]
---*/

assert.sameValue(typeof Float32Array, 'function', 'typeof Float32Array is "function"');
