<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <title>grid-template-rows neutral keyframe</title>
    <link rel="help" href="https://drafts.csswg.org/css-grid/#track-sizing">
    <meta name="assert" content="grid-template-rows supports neutral keyframe.">
    <script src="/resources/testharness.js"></script>
    <script src="/resources/testharnessreport.js"></script>
    <script src="/css/support/interpolation-testcommon.js"></script>
  </head>
  <body>
    <div id="target"></div>
    <script>
      test(() => {
        target.style.gridTemplateRows = '200px';
        var animation = target.animate([{}, {gridTemplateRows: '30px 300px'}], 1000);
        animation.pause();
        animation.currentTime = 500;
        assert_equals(getComputedStyle(target).gridTemplateRows, '30px 300px');
        target.style.gridTemplateRows = '20px 200px';
        assert_equals(getComputedStyle(target).gridTemplateRows, '25px 250px');
        target.style.gridTemplateRows = '200px';
        assert_equals(getComputedStyle(target).gridTemplateRows, '30px 300px');
      }, 'Invalidate the result for the neutral keyframe if the size of the list changes.');
    </script>
  </body>
</html>