#!

// Copyright (C) 2019 Leo Balter. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: pending
description: >
    Hashbang comments should be allowed in Modules.
info: |
    HashbangComment::
      #! SingleLineCommentChars[opt]
flags: [module, raw]
features: [hashbang]
---*/
