<!doctype html>
<html>
    <head>
        <title>HTML5 Drag and Drop: Auto state of draggable attribute for 'a' element</title>
        <meta content="text/html; charset=UTF-8" http-equiv="Content-Type"/>
        <link rel="author" title="Microsoft" href="http://www.microsoft.com/"/>
        <link rel="help" href="http://dev.w3.org/html5/spec/dnd.html#the-draggable-attribute"/>
        <meta name="assert" content="Auto state of draggable attribute for 'a' element"/>
        <script src="/resources/testharness.js"></script>
        <script src="/resources/testharnessreport.js"></script>
    </head>
    <body>
        <div id=log></div>
        <a href="#" style="display:none" id="target">Drag this</a>
        <script type="text/javascript">
            test(function() {assert_true(document.getElementById("target").draggable)}, "Auto state of draggable attribute for 'a' element");
        </script>
    </body>
</html>
