// Copyright (C) 2015 André Bargull. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-typedarray-constructors
description: >
  Float64Array is a constructor function.
features: [TypedArray]
---*/

assert.sameValue(typeof Float64Array, 'function', 'typeof Float64Array is "function"');
