<!DOCTYPE html>
<meta charset="utf-8">
<title>HTML Test: focusgroup - none behavior token disables navigation</title>
<meta name="assert" content="none disables all focusgroup navigation.">
<link rel="author" title="Microsoft" href="http://www.microsoft.com/">
<link rel="help" href="https://open-ui.org/components/scoped-focusgroup.explainer/">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src="/resources/testdriver.js"></script>
<script src="/resources/testdriver-vendor.js"></script>
<script src="/resources/testdriver-actions.js"></script>
<script src="/shadow-dom/focus-navigation/resources/focus-utils.js"></script>
<script src="../resources/focusgroup-utils.js"></script>

<div focusgroup="none">
  <button id=n1>n1</button>
  <button id=n2>n2</button>
</div>

<script>

  promise_test(async t => {
    const n1 = document.getElementById('n1');
    await focusAndKeyPress(n1, kArrowRight);
    assert_equals(document.activeElement.id, 'n1');
  }, "none: ArrowRight blocked");

  promise_test(async t => {
    const n1 = document.getElementById('n1');
    await focusAndKeyPress(n1, kArrowLeft);
    assert_equals(document.activeElement.id, 'n1');
  }, "none: ArrowLeft blocked");

  promise_test(async t => {
    const n1 = document.getElementById('n1');
    await focusAndKeyPress(n1, kArrowDown);
    assert_equals(document.activeElement.id, 'n1');
  }, "none: ArrowDown blocked");

  promise_test(async t => {
    const n1 = document.getElementById('n1');
    await focusAndKeyPress(n1, kArrowUp);
    assert_equals(document.activeElement.id, 'n1');
  }, "none: ArrowUp blocked");

</script>
