// Copyright (C) 2017 Leo Balter. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-array-constructor-array
description: >
  The Array constructor is a built-in function
---*/

assert.sameValue(typeof Array, 'function', 'The value of `typeof Array` is expected to be "function"');
