// Copyright (C) 2015 André Bargull. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-typedarray-constructors
description: >
  Uint16Array is a constructor function.
features: [TypedArray]
---*/

assert.sameValue(typeof Uint16Array, 'function', 'typeof Uint16Array is "function"');
