// Copyright (C) 2017 André Bargull. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-identifiers-static-semantics-early-errors
description: >
  `yield` is not a reserved identifier in non-strict mode code and may be used as a label.
flags: [noStrict]
---*/

yi\u0065ld: 1;
