// Copyright 2009 the Sputnik authors.  All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-future-reserved-words
es5id: 7.6.1.2_A1.8
description: Checking if execution of "double=1" succeeds
info: |
  `double` was removed from the list of future reserved words in ECMAScript 5.
---*/

var double = 1;
