// Copyright 2009 the Sputnik authors.  All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-future-reserved-words
es5id: 7.6.1.2_A1.14
description: Checking if execution of "goto=1" succeeds
info: |
  `goto` was removed from the list of future reserved words in ECMAScript 5.
---*/

var goto = 1;
