// Copyright (C) 2015 André Bargull. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-typedarray-constructors
description: >
  Int16Array is a constructor function.
features: [TypedArray]
---*/

assert.sameValue(typeof Int16Array, 'function', 'typeof Int16Array is "function"');
