// Copyright (C) 2015 André Bargull. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-typedarray-constructors
description: >
  Int32Array is a constructor function.
features: [TypedArray]
---*/

assert.sameValue(typeof Int32Array, 'function', 'typeof Int32Array is "function"');
