// Copyright (C) 2021 Microsoft. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-array.prototype.findlast
description: >
  Return abrupt from predicate call.
info: |
  Array.prototype.findLast ( predicate[ , thisArg ] )

  ...
  4. Let k be len - 1.
  5. Repeat, while k â‰¥ 0,
    ...
    c. Let testResult be ! ToBoolean(? Call(predicate, thisArg, Â« kValue, ğ”½(k), O Â»)).
    d. If testResult is true, return kValue.
  ...
features: [array-find-from-last]
---*/

var predicate = function() {
  throw new Test262Error();
};

assert.throws(Test262Error, function() {
  [1].findLast(predicate);
});
